---
title: Course Materials
---

## Textbooks

Reading is important in learning a new language. This is especially so for learning a programming language, whose syntax has a very precise structure and interpretation. Each topic has some required readings from the textbook [Halterman, Richard L. "Fundamentals of Python programming." (2018)](https://archive.org/details/2018Fundamentals.ofPython). As the programming languages evolve over times, you may supplement your readings with the Python documentations for the later versions:

- [Python 3.12](https://docs.python.org/3.12/) used for quiz and exam, and
- [Python 3.11](https://docs.python.org/3.11/) used for lab assignents.

In comparison, the textbook uses a much earlier version [Python 3.4](https://docs.python.org/3.4/).

::::{caution} Is it okay to learn/use different versions?

The version number such as `3.12` consists of the major version number `3` followed by the minor version number `12`.

- It is okay to learn the earlier version `3.4` as Python's design has made a conscious effort to maintain backwards compatibility for different minor versions, i.e., code written in an earlier minor version should run in a later minor version.
- Version `3.12` is used for quiz and exam as it is the latest stable version, which is often more efficient.
- Version `3.11` is used for lab assignments because the assignments often require external libraries that may not yet support Python's latest stable version.

However, new syntax in later versions may not work in earlier versions.

::::

As a student at CityU, you also have access to a wide range of library resources to supplement your reading. A useful e-learning resource to help you dive deeper into the cutting-edge computing technologies is LinkedIn Learning:

1. Open the [LinkedIn Learning](https://www.cityu.edu.hk/its/services-facilities/online-courses-linkedin-learning) page.
2. Click <kbd>LinkedIn Learning Login</kbd> to login with your CityU account.

There is a vast collection of online courses and tutorials on various topics. For instance:

::::{seealso} LinkedIn Learning
:class: dropdown

Some courses on introductory python programming:

- [Python Essential Training](https://www.linkedin.com/learning-login/share?account=76816450&forceAccount=false&redirect=https%3A%2F%2Fwww.linkedin.com%2Flearning%2Fpython-essential-training-18764650%3Ftrk%3Dshare_ent_url%26shareId%3DaXCm1ZkcTTCExEgNCUlh%252Bw%253D%253D)
- [Learning the Python 3 Standard Library](https://www.linkedin.com/learning-login/share?account=76816450&forceAccount=false&redirect=https%3A%2F%2Fwww.linkedin.com%2Flearning%2Flearning-the-python-3-standard-library%3Ftrk%3Dshare_ent_url%26shareId%3DGcv4DyTKSuWBBsKpEmIamQ%253D%253D)


A course on how to use LinkedIn Learning: 

- [How to Use LinkedIn Learning](https://www.linkedin.com/learning-login/share?account=76816450&forceAccount=false&redirect=https%3A%2F%2Fwww.linkedin.com%2Flearning%2Fhow-to-use-linkedin-learning%3Ftrk%3Dshare_ent_url%26shareId%3DT0Oya4yGTOGygXyXbWpMCQ%253D%253D)

::::

The following are similar introductory python programming courses at other universities:[^1]

[^1]: Note that the syllabi have different coverages even though they both assume the students have no prior programming experience.

::::{seealso} MIT Open Courseware
:class: dropdown

An introductory programming course using Python:

- Course homepage: [6.0001
Introduction To Computer Science And Programming In Python](https://ocw.mit.edu/courses/6-0001-introduction-to-computer-science-and-programming-in-python-fall-2016/pages/readings/)  
  Textbook: [Guttag, John V. Introduction to computation and programming using Python: with application to computational modeling and understanding data. Mit Press, 2021.](https://julac-cuh.primo.exlibrisgroup.com/permalink/852JULAC_CUH/vit3jk/alma991029721214003408)

An introductory programming course using Scheme:

- Course homepage: [6.001 Structure And Interpretation Of Computer Programs](https://ocw.mit.edu/courses/6-001-structure-and-interpretation-of-computer-programs-spring-2005/)
  Textbook: [Abelson, Harold, and Gerald Jay Sussman. Structure and interpretation of computer programs. The MIT Press, 1996.](https://web.mit.edu/6.001/6.037/sicp.pdf)


::::

::::{seealso} UC Berkeley
:class: dropdown

An introductory programming course using Python:

- Course homepage: [CS 61A Structure And Interpretation Of Computer Programs](https://cs61a.org)  
  Textbook: [John DeNero. Composing Programs.](https://www.composingprograms.com/)


::::

## Notebooks

Lecture and lab materials are written in the form of Jupyter notebooks. Jupyter notebooks provide an interactive <wiki:Literate_programming> experience, which is a significant benefit for learners:

- Jupyter Notebooks allow programs to be mixed with other contents, including figures, videos, and formatted textual explanations.
- The notebooks can also run in a Jupyter server, which enables users to write and run their programs while adding formatted notes to it.

To fetch the notebooks, follow the links to the notebooks on the course homepage such as:

::::{card}
:header: An example of a notebook link on the course homepage.

<div style="border: 1px solid black; padding: 10px;">
<h4 title="Title: Lecture/Lab">Lab</h4>
<ul>
  <li>
    <a title="Subtitle: Notebook link" href="https://canvas.cityu.edu.hk/courses/48496/external_tools/retrieve?display=borderless&url=https%3A//dive.cs.cityu.edu.hk/cs1302_24a/hub/lti/launch%3Fcustom_next%3Dhttps%253A//dive.cs.cityu.edu.hk/cs1302_24a/hub/user-redirect/git-pull%253Frepo%253Dhttps%25253A//github.com/dive4dec/cs1302_24a%2526urlpath%253Dlab/tree/cs1302_24a/Lab0/Course_Materials.ipynb">Course Materials</a>
</ul>
</div>

After clicking the link above, you may be asked to login with your CityU account and specify a server option. Select the <kbd>Default</kbd> option from the list of available Jupyter servers, and click <kbd>Start</kbd> to begin your session.

::::

::::{note} What are the server options?

The course's programming environment is conveniently accessible via a [JupyterHub][jh] server, allowing remote access without the need to install any special software, thanks to <wiki:Project_Jupyter>. This means you can easily write and run programs on your mobile devices using just a web browser. Each student will have their own [Jupyter server](https://jupyter-server.readthedocs.io/en/latest/) with individual computer resources such as CPU/memory/GPU.

- Each server runs the <wiki:Ubuntu> operating system in a [container](wiki:Containerization_(computing)). See the [Dockerfile](https://github.com/dive4dec/jupyter/blob/main/cs1302nb/Dockerfile) for details.
- The GPU resources are useful for customized AI applications.
- You can switch between the `Default` and other server options at any time by restarting your server and selecting the desired option. Restarting the server is simple and can be done through the hub control panel, as described in the troubleshooting note below.

::::

[jh]: https://jupyterhub.readthedocs.io/en/stable/

If the server is spawned successfully, the [JupyterLab](https://jupyterlab.readthedocs.io) interface should appear.

:::{tip} Troubleshooting[^click]
:class: dropdown

In case the jupyterlab interface fails to show:

- Refresh the page or restart your browser.
- Restart your server using the `Hub Control Panel` under the `File` and click `Stop My Server` and then `Start My Server`.

:::

[^click]:  click me to expand

As a demonstration of the benefit of literate programming, run the following code to start the notebook link generator app.[^run]

[^run]: To run the cell, select it and press <kbd>Shit + Enter</kbd>.

```python
from notebook_link_generator import setup_notebook_link_widget

display_widget = setup_notebook_link_widget()
display(display_widget)
```

Change the `Notebook path` above to `Lecture1/Introduction_to_Computer_Programming.ipynb` and observe the live update to the generated link.

Run the following code to see how the notebook was generated steyp-by-step:[^inspect]

[^inspect]: You can inspect parts of the code by placing the cursor there and press <kbd>Shift + Tab</kbd>.

```python
from urllib.parse import quote

# Course parameters
course_id = "cs1302_24a"
notebook_path = "Lecture1/Introduction_to_Computer_Programming.ipynb"
canvas_id = "48496"

# Base URLs
notebook_repo = f"https://github.com/dive4dec/{course_id}"
course_server = "https://dive.cs.cityu.edu.hk"
course_homepage = f"https://canvas.cityu.edu.hk/courses/{canvas_id}"

# Notebook URL generation
# 1. Construct the URL for git-pull service to clone or pull the notebook repository
gitpull_url = (
    f"{course_server}/{course_id}/hub/user-redirect/git-pull?"
    + f"repo={quote(notebook_repo)}"
    + f"&urlpath={quote(f'lab/tree/{course_id}/{notebook_path}')}"
)
# Construct the LTI (Learning Tools Interoperability) launch URL for JupyterHub
lti_url = (
    f"{course_server}/{course_id}/hub/lti/launch?" + f"custom_next={quote(gitpull_url)}"
)
# Construct the LTI external tool URL to open the notebook
notebook_url = (
    f"{course_homepage}/external_tools/retrieve?display=borderless&"
    + f"url={quote(lti_url)}"
)
notebook_url
```

::::{note} What does the notebook link do?

The notebooks reside in a [GitHub repository](https://github.com/dive4dec/cs1302_24a/blob/main/README.md). Accessing a notebook link will

1. Uses [LTI Authenticator](https://ltiauthenticator.readthedocs.io) to launch the JupyterHub as an LTI external tool from the course homepage.
2. Uses [NbGitPuller](https://jupyterhub.github.io/nbgitpuller/index.html) to [git-pull](https://git-scm.com/docs/git-pull)s all the files in the repository, merge them with any existing files stored under the course folder `cs1302_24a` in your home directory without overwriting your changes, and open the notebook path specified in the JupyterLab interface.

::::

## Assignments

### How to complete?

The notebooks can be edited in JupyterLab to include your answers for submission. If this is your first time using JupyterLab, take a look at the [official video tutorial](https://www.youtube.com/embed/A5YyoCKxEOU):

::::{seealso} How to use JupyterLab?
:class: dropdown

:::{iframe} https://www.youtube.com/embed/A5YyoCKxEOU
:width: 100%
:align: left

Official video tutorial on Jupyter
:::
::::

For more advanced features:

- Checkout the `Help` menu items
    - `Show Keyboard Shortcuts` to try some of the shortcuts to see their effect; and
    -  `Jupyter Reference` to open the [user guide](https://jupyterlab.readthedocs.io/en/latest/user/interface.html) in a new tab.
- Try also the [MyST Markdown](https://mystmd.org/guide/typography) syntax to format your notes.

In learning a new programming language, the first program to write is often the ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program, which says Hello to the world. As your first lab exercise, you will write such a program in python.

:::::{exercise} Hello World
:label: ex:Hello-World

Complete the program to print the message `Hello, World!`. 

::::{hint}
:class: dropdown

One possible solution is:

:::{code-block} python
:name: code:Hello-World
:caption: Python, "Hello, World!"

def say_hello():
    print("Hello, World!")
    
say_hello()
:::

:::{caution}
You must use a tab or 4 spaces to indent the second line of code `print(...)`.
:::

::::

:::::

The following code cell is a [solution cell](https://nbgrader.readthedocs.io/en/stable/configuration/student_version.html#default-behavior). Since you are not expected to know python yet, you can simply expand the hint above and copy the answer to the solution cell instead.

<!-- grade_id: helloworld, points: 0, grade: false, solution: true -->
```python
def say_hello():
    # YOUR CODE HERE
    raise NotImplementedError()
    
say_hello()
```

:::{attention}

It's important to follow certain guidelines when writing your answers or code in the notebooks:

- Only provide your answers in the cells that are designated for this purpose, which are usually labeled with `YOUR CODE HERE` or `YOUR ANSWER HERE`.
- For coding exercises, be sure to remove the line `raise NotImplementedError()` from the cell before submitting your work.
- Do not clone or duplicate the answer cells, as this can cause issues with version control and grading.

:::

To test your program:

- Run your program by selecting the solution cell and press `Shift`+`Enter`.
- Then, run the following visible test to check whether your program prints the correct message.

<!-- grade_id: test-helloworld, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your "Hello, World!" program.
import io
import sys

old_stdout, sys.stdout = sys.stdout, io.StringIO()
say_hello()
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == "Hello, World!\n"
```

The test returns an assertion error if your program does not print the correct message.

::::{tip}

- You can repeatedly modify your solution and run the test cell until your solution passes the test. There is no mark penalty in failing the test before submission. 
- You should try to understand the error messages from a failed test.
- To assess your solution thoroughly, we will run new tests hidden from you after you have submitted your notebook. Therefore, *you should ensure your solution works in general rather than just the visible tests*.
- If you open the same notebook multiple times in different browser windows, be careful in making changes in different windows as you may overwrite your own changes.
- If your notebook fails to run any code, the kernel might have died. You can restart the kernel with `Kernel`$\to$ `Restart`. If restarting fails, check your code cells to see if there is anything that breaks the kernel.

::::

### How to submit?

:::{important} Grading policies

- Late submissions will not be accepted without valid justifications.
- You can submit your assignment repeatedly before the deadline without penalty, so it is recommended to submit your assignment in advance. Please note that you are responsible for any issues related to failed submissions due to network outages that occur too close to the deadline.
- It is also important to make a record or backup of your submission that includes the submission timestamp. Double check to ensure that you have submitted the correct lab assignment, since multiple lab assignments may be open for submission at the same time.

:::

You normally have at least 5 days to work on the lab after your lab session. 

- You can check the due dates of all the labs from the course homepage. 
- You may seek help from us or your classmates. However, you must write your own solution and indicate who your collaborators are by including their EIDs in a code cell:  
    ```Python
    COLLABORATORS = ["xfwong2", "mklee3"]
    ```

:::{attention}

- Lecture notebooks under the subfolders `Lecture*\` need NOT be submitted. You only need to submit the lab notebooks under `Lab*\`. 

:::

Before you submit, make sure everything runs as expected:

1. **Git-pull the notebooks** (optional): Follow any one of the link on the course homepage to a notebook, which will git-pull any updates/corrections to (all) your notebooks.
1. **Save the notebook**: Unsaved changes are not exported, even though they are in the memory.
1. **Restart the kernel**: `Kernel`$\to$ `Restart Kernel...` to have a clean state before running visible tests.
1. **run all cells**: `Run`$\to$ `Run All Cells` to double check the results of the visible tests are as expected.

To submit your notebook:

1. Select the menu item `Nbgrader`$\to$`Assignment List`.
1. Expand the Lab folder and click the `validate` button next to the notebook(s) to check if all the visible tests pass.
1. Click `Submit` to submit your notebook.

::::{seealso} What is NBGrader?
:class: dropdown

[Nbgrader](https://nbgrader.readthedocs.io/en/stable/) is a package for grading notebook assignments. It is well-integrated into the JupyterLab interface, allowing students to submit their notebooks directly through the JupyterHub server. Submitted notebooks can be both auto-graded with pre-defined test cases and manually graded with custom feedback. After grading is complete, you can check your scores and access the feedback using the same interface.

::::

::::{caution}

Your submission may not be graded under the following circumstances:

- If the notebook files have been renamed, for example, `Setup.ipynb` being changed or copied to `Setup (Copy).ipynb`.
- If an HTML file exists with the same name as a notebook, for example, `Setup.html`.
- If the file size exceeds `100MB`.
- If the code takes too long to run or requires an excessive amount of memory.

It is essential to ensure that your submission meets these guidelines to avoid any issues with grading.

:::{tip} Troubleshooting
:class: dropdown

If you believe your notebooks are corrupted, 
1. download/backup your existing notebooks,
1. remove them from the `Lab*/` folder,
1. click the git-pull links from the course homepage to re-pull the notebooks, and
1. copy your solutions to the new notebooks.

You may also run the course notebooks outside the jupyterhub server and locally on your computer. For more details, see the course homepage.

:::

::::