---
title: Binary Code
---

Welcome to this lab where you will explore how computers represent numbers using a fun card-guessing game! Through this lab, you will gain an understanding of how binary numbers are used to represent decimal numbers in computers.

## Card Guessing Game

### Rules

::::{list-table} Deck of cards
:name: tab:deck
:header-rows: 1

*   -
    - 1
    - 2
    - 3
    - 4
    - 5
    - 6
    - 7
    - 8
    - 9
    - 10
    - 11
    - 12
    - 13
*   - Diamond
    - ![A-Diamond](https://upload.wikimedia.org/wikipedia/commons/e/e9/Cards-A-Diamond.svg)
    - ![2-Diamond](https://upload.wikimedia.org/wikipedia/commons/9/99/Cards-2-Diamond.svg)
    - ![3-Diamond](https://upload.wikimedia.org/wikipedia/commons/4/44/Cards-3-Diamond.svg)
    - ![4-Diamond](https://upload.wikimedia.org/wikipedia/commons/a/af/Cards-4-Diamond.svg)
    - ![5-Diamond](https://upload.wikimedia.org/wikipedia/commons/d/dd/Cards-5-Diamond.svg)
    - ![6-Diamond](https://upload.wikimedia.org/wikipedia/commons/4/44/Cards-6-Diamond.svg)
    - ![7-Diamond](https://upload.wikimedia.org/wikipedia/commons/2/2b/Cards-7-Diamond.svg)
    - ![8-Diamond](https://upload.wikimedia.org/wikipedia/commons/9/90/Cards-8-Diamond.svg)
    - ![9-Diamond](https://upload.wikimedia.org/wikipedia/commons/2/25/Cards-9-Diamond.svg)
    - ![10-Diamond](https://upload.wikimedia.org/wikipedia/commons/c/c2/Cards-10-Diamond.svg)
    - ![J-Diamond](https://upload.wikimedia.org/wikipedia/commons/7/78/Cards-J-Diamond.svg)
    - ![Q-Diamond](https://upload.wikimedia.org/wikipedia/commons/c/c3/Cards-Q-Diamond.svg)
    - ![K-Diamond](https://upload.wikimedia.org/wikipedia/commons/5/55/Cards-K-Diamond.svg)
*   - Club
    - ![A-Club](https://upload.wikimedia.org/wikipedia/commons/c/c4/Cards-A-Club.svg)
    - ![2-Club](https://upload.wikimedia.org/wikipedia/commons/b/b0/Cards-2-Club.svg)
    - ![3-Club](https://upload.wikimedia.org/wikipedia/commons/e/e0/Cards-3-Club.svg)
    - ![4-Club](https://upload.wikimedia.org/wikipedia/commons/6/69/Cards-4-Club.svg)
    - ![5-Club](https://upload.wikimedia.org/wikipedia/commons/7/7e/Cards-5-Club.svg)
    - ![6-Club](https://upload.wikimedia.org/wikipedia/commons/a/af/Cards-6-Club.svg)
    - ![7-Club](https://upload.wikimedia.org/wikipedia/commons/8/8e/Cards-7-Club.svg)
    - ![8-Club](https://upload.wikimedia.org/wikipedia/commons/f/fd/Cards-8-Club.svg)
    - ![9-Club](https://upload.wikimedia.org/wikipedia/commons/a/ac/Cards-9-Club.svg)
    - ![10-Club](https://upload.wikimedia.org/wikipedia/commons/2/25/Cards-10-Club.svg)
    - ![J-Club](https://upload.wikimedia.org/wikipedia/commons/c/c7/Cards-J-Club.svg)
    - ![Q-Club](https://upload.wikimedia.org/wikipedia/commons/3/37/Cards-Q-Club.svg)
    - ![K-Club](https://upload.wikimedia.org/wikipedia/commons/9/9e/Cards-K-Club.svg)
*   - Heart
    - ![A-Heart](https://upload.wikimedia.org/wikipedia/commons/6/60/Cards-A-Heart.svg)
    - ![2-Heart](https://upload.wikimedia.org/wikipedia/commons/6/6e/Cards-2-Heart.svg)
    - ![3-Heart](https://upload.wikimedia.org/wikipedia/commons/5/57/Cards-3-Heart.svg)
    - ![4-Heart](https://upload.wikimedia.org/wikipedia/commons/3/39/Cards-4-Heart.svg)
    - ![5-Heart](https://upload.wikimedia.org/wikipedia/commons/9/91/Cards-5-Heart.svg)
    - ![6-Heart](https://upload.wikimedia.org/wikipedia/commons/5/55/Cards-6-Heart.svg)
    - ![7-Heart](https://upload.wikimedia.org/wikipedia/commons/d/d4/Cards-7-Heart.svg)
    - ![8-Heart](https://upload.wikimedia.org/wikipedia/commons/5/55/Cards-8-Heart.svg)
    - ![9-Heart](https://upload.wikimedia.org/wikipedia/commons/d/d2/Cards-9-Heart.svg)
    - ![10-Heart](https://upload.wikimedia.org/wikipedia/commons/7/76/Cards-10-Heart.svg)
    - ![J-Heart](https://upload.wikimedia.org/wikipedia/commons/e/e7/Cards-J-Heart.svg)
    - ![Q-Heart](https://upload.wikimedia.org/wikipedia/commons/2/28/Cards-Q-Heart.svg)
    - ![K-Heart](https://upload.wikimedia.org/wikipedia/commons/b/bf/Cards-K-Heart.svg)
*   - Spade
    - ![A-Spade](https://upload.wikimedia.org/wikipedia/commons/9/9d/Cards-A-Spade.svg)
    - ![2-Spade](https://upload.wikimedia.org/wikipedia/commons/e/e7/Cards-2-Spade.svg)
    - ![3-Spade](https://upload.wikimedia.org/wikipedia/commons/d/d0/Cards-3-Spade.svg)
    - ![4-Spade](https://upload.wikimedia.org/wikipedia/commons/4/4e/Cards-4-Spade.svg)
    - ![5-Spade](https://upload.wikimedia.org/wikipedia/commons/b/b1/Cards-5-Spade.svg)
    - ![6-Spade](https://upload.wikimedia.org/wikipedia/commons/6/68/Cards-6-Spade.svg)
    - ![7-Spade](https://upload.wikimedia.org/wikipedia/commons/c/c6/Cards-7-Spade.svg)
    - ![8-Spade](https://upload.wikimedia.org/wikipedia/commons/7/7e/Cards-8-Spade.svg)
    - ![9-Spade](https://upload.wikimedia.org/wikipedia/commons/0/0a/Cards-9-Spade.svg)
    - ![10-Spade](https://upload.wikimedia.org/wikipedia/commons/6/67/Cards-10-Spade.svg)
    - ![J-Spade](https://upload.wikimedia.org/wikipedia/commons/e/ea/Cards-J-Spade.svg)
    - ![Q-Spade](https://upload.wikimedia.org/wikipedia/commons/e/ef/Cards-Q-Spade.svg)
    - ![K-Spade](https://upload.wikimedia.org/wikipedia/commons/1/18/Cards-K-Spade.svg)

::::

[](#tab:deck) shows a deck of 52 cards:
-   Each card is in one of the four suits: *Diamond* &#9830;, *Club* &#9827;, *Heart* &#9829;, and *Spade* &#9824;.
-   Each card has a value ranging from 1 to 13. For simplicity, the original card values A, J, Q, and K are converted to the numerical values 1, 11, 12, and 13 respectively.

:::{figure} images/card_game.dio.svg
:name: fig:card-guessing
:alt: A player guesses a card drawn by a dealer.
:align: left

Card guessing game
:::

::::{card}
:header: Rules

As depicted in [](#fig:card-guessing), the card-guessing game involves the following steps:

- The dealer selects a card without revealing it to the player.
- The player's goal is to correctly guess the chosen card.
- To make an informed guess, the player can ask up to six yes/no questions.
- The dealer must answer each question truthfully and immediately.

:::{prf:example} Valid questions
:class: dropdown
:label: eg:valid-questions

For instance, the player may ask:
- Is the suit club?
- Is the card diamond 1?
- Is the value at least 10?

:::

:::{prf:example} Invalid questions
:class: dropdown
:label: eg:invalid-questions

However, the player cannot ask:  
- What is the value?
- What is the suite?

because the answers are not **yes/no**.

:::

::::

One strategy is to ask whether the randomly drawn card is a specific card, e.g.,

:::{card}

Is it the Diamond 1?

:::

If the answer is yes, the player can confidently guess that the card is Diamond 1 and win the game. However, if the answer is no, the player can continue to check another card, e.g.,

:::{card}

Is it the Diamond 2?

:::

and so on...

A {term}`pseudocode` that summarizes the above strategy (or {term}`algorithm`) is shown below.

::::{prf:algorithm} Checking one card at a time 
:label: code:1-by-1

```{code} python
:linenos:
:emphasize-lines: 3,4

while the number of questions asked < 6 and \
      the correct card has not been guessed yet:
    randomly pick a new card not chosen before
    ask the dealer whether the drawn card is the chosen one
    if the answer is yes:
        guess the card
    else:
        increment the number of questions asked
```

::::

::::{exercise} chance
:label: ex:chance

What is the chance of winning using the strategy in [](#code:1-by-1)? 

:::{attention}
Provide your answer in the following code cell by assigning a fraction to the variable `chance` such as

```python
chance = 0.5
```

where the value 0.5 means 50% chance. This is just a demo, i.e., 50% is *not* necessarily the correct answer. Your answer should be accurate at least up to *2 decimal places*.

:::

::::

<!-- grade_id: chance, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError
```

::::{tip}

Run the following visible tests to check whether the `chance` variable meets the following criteria:

- It is greater than or equal to 0.
- It is less than or equal to 1.

::::

<!-- grade_id: test-chance, points: 1, grade: true, solution: false -->
```python
# tests
assert chance >= 0
assert chance <= 1
```

::::{caution}

Passing the visible tests does not guarantee that the answer is correct. The actual correctness will be evaluated using a hidden test to be injected into the following cell.

::::

<!-- grade_id: htest-chance, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

### Virtual Cards

Instead of drawing a physical card, the dealer can use programming to generate a virtual deck and draw a *random*[^pseudorandom] virtual card from it. Run the code below to get the required tools.

[^pseudorandom]: The truth is, the random generation process will look random but is not truly random. You will learn about this pseudorandomness later in the course.

```python
from collections import namedtuple  # for naming the cards
from random import choice  # for randomly drawing cards
```

You'll learn the syntax in a future lecture, but even now, it is clear that the code imports {term}`libraries <library>` for creating named tuples and making random choices, thanks to the design of Python being a high-level programming language.

To create a `deck` of cards, run the following cell.

```python
suits = ("Diamond", "Club", "Heart", "Spade")
values = range(1, 14)
Card = namedtuple("Card", ["value", "suit"])
deck = [Card(value, suit) for value in values for suit in suits]
deck
```

The code above uses {term}`composite data types <composite data type>` like tuples and lists, which will be formally introduced later in the course. They are handy for storing and manipulating collections of values.

Now, you can use the following program to play the game with your friends:

```python
choice(deck)
```

::::{tip}

To draw cards repeatedly, use <kbd>Ctrl-Enter</kbd> instead of <kbd>Shift-Enter</kbd> to run the code, so the same code cell stay selected after the execution. Note that function `choice` is said to be an [impure function](https://en.wikipedia.org/wiki/Pure_function) as it can return different values for repeated runs with the same input argument, namely `deck`.

::::

::::{exercise} Draw multiple cards
:label: ex:multiple-draw

Can you fix the code below so that it can return 5 cards randomly drawn with replacement from the deck of cards?

```{code} python
:linenos:
:emphasize-lines: 2

def draw_five_cards_randomly():
    card1 = card2 = card3 = card4 = card5 = choice(deck)
    return  card1, card2, card3, card4, card5
```

::::

<!-- grade_id: multiple_draw, points: 0, grade: false, solution: true -->
```python
def draw_five_cards_randomly():
    # YOUR CODE HERE
    raise NotImplementedError
    return  card1, card2, card3, card4, card5
```

<!-- grade_id: test-multiple_draw, points: 1, grade: true, solution: false -->
```python
# tests
import random

random.seed(1302)  # to make pseudo-random sequences reproducible
assert (lambda cards: any(cards[i] != cards[4] for i in range(4)))(
    draw_five_cards_randomly()
)
```

<!-- grade_id: htest-multiple_draw, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

::::{caution}
:class: dropdown

While one may pass the visible test with an answer like

```
def draw_five_cards_randomly():
    card1 = card2 = card3 = card4 = card5 = choice(deck)
    card1 = 0
    return  card1, card2, card3, card4, card5
```

Who knows whether the hidden test would catch that `card1` is not even a card in the `deck`? The hidden test can also be modified since it is not hard coded in the notebook.

::::

### Virtual Player

Given you have randomly drawn a card,

1. run the following code cell, and
2. answer the 6 questions raised by the player honestly by entering  
    - `1` for yes and
    - `0` for no.

```python
def is_yes(question):
    """Get an answer to a yes/no question."""
    return "1" == input(question + " 1:yes [0:no] ")


suit_idx = number = 0
if is_yes("Is the suite either heart or spade?"):
    suit_idx += 2
if is_yes("Is the suite either club or spade?"):
    suit_idx += 1
for i in range(3, -1, -1):
    number += 2**i * is_yes(f"Is the number {number+2**i} or above?")
print(f"The card is {suits[suit_idx]} {number}")
```

The code above provides an overview of various programming syntaxes to be covered in the course. Although you don't need to be familiar with them now, it's good to have a basic understanding:

- The program first defines a function called `is_yes` that checks whether the answer to a given question is yes. 
- To compute the guess, the program modifies the values of the variables `suit_idx` and `number` using
  - {term}`conditional` `if` statements and
  - a `for` {term}`loop <iteration>`, along with
  - different operations such as
    - chained assignment `=`,
    - augmented assignment `+=`,
    - equality comparison `==`, and
    - exponentiation `**`.

::::{exercise} Resolution of uncertainty
:label: ex:resolution-capacity

Hope you are convinced that the virtual player has a winning strategy. Indeed, the answers to 6 questions not only resolve the uncertainty of a random draw from 6 cards, but they can also resolve a random draw from over 52 cards. Fix the following function number_of_cards so that it returns the maximum number of cards from which a random draw can be resolved by the answers to a number of questions specified by an integer assigned to the variable `number_of_questions`.

```{code} python
:linenos:
:emphasize-lines: 2

def number_of_cards(number_of_questions):
    ans = number_of_questions
    return ans
```

You may assume one can add new values/suits so that a deck may have an arbitrarily number of distinct cards.

::::

<!-- grade_id: resolution_capacity, points: 0, grade: false, solution: true -->
```python
def number_of_cards(number_of_questions):
    # YOUR CODE HERE
    raise NotImplementedError
    return ans
```

<!-- grade_id: test-resolution_capacity, points: 1, grade: true, solution: false -->
```python
# tests
assert number_of_cards(number_of_questions=6) > len(deck) 
```

<!-- grade_id: htest-resolution_capacity, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

::::{caution}

The value of `number_of_questions` can be very large in the hidden test.

::::

## Binary Number System

### Representing non-negative integers

The following table gives the binary representions of unsigned decimal integers from 0 to 7.

::::{list-table} Encoding positive integers
:name: tab:pos-int
:header-rows: 1

*
    - Binary
    - Decimal
*
    - 000
    - 0
*
    - 001
    - 1
*
    - 010
    - 2
*
    - 011
    - 3
*
    - **1**00
    - **4**
*
    - **1**01
    - **5**
*
    - **1**10
    - **6**
*
    - **1**11
    - **7**

::::

Observe that the binary representations of 4, 5, 6, and 7 all have **1** in the leftmost (most significant) bit. Consider that bit as the answer to the following **yes/no** question when converting an integer `n` to binary:

> Is the integer 4 or above, i.e., $n\geq 4$?

Now, to convert the entire binary sequence to decimal, we can think of the conversion as a guessing game where
- the binary sequence is a sequence of **yes (1)** or **no (0)** answers to certain **yes/no** questions, and
- the informed guess is the integer represented by the binary sequence.

![Decoding tree](images/dt.dio.svg)

The following code attempts to convert integers between 0 and 7 to their binary strings.

```python
def decimal_to_binary(n):
    def get_bit(cond):
        return "1" if cond else "0"
    return get_bit(eval(Q1)) + get_bit(eval(Q2)) + get_bit(eval(Q3))

# Define the questions
Q1, Q2, Q3 = "n >= 4", "n >= 2", "n >= 1"

# Print the conversions decimal (binary)
print(*[f"{i} ({decimal_to_binary(i)})" for i in range(8)], sep="\n")
```

:::::{exercise} question to ask
:label: ex:d2b

Although the conversions for 0, 1, 3, and 7 are correct, the conversions for 2, 4, 5, and 6 are incorrect. Fix the issue by assigning the correct strings to Q2 and Q3. Similar to Q1, each of the strings should be a valid Python expression that can be evaluated with `n` being the integer to convert to a binary string. 

::::{hint}
:class: dropdown

Consider the modulo operator `%`.

:::{seealso}
- [Binary Number System](https://www.mathsisfun.com/binary-number-system.html) and
- [Binary Number Conversions](https://www.purplemath.com/modules/numbbase.htm).
:::
::::


:::::

<!-- grade_id: d2b, points: 0, grade: false, solution: true -->
```python
# Assign strings to Q2 and Q3
# YOUR CODE HERE
raise NotImplementedError
Q2, Q3
```

<!-- grade_id: test-d2b, points: 1, grade: true, solution: false -->
```python
# test
assert decimal_to_binary(2) == '010'
assert decimal_to_binary(4) == '100'
assert decimal_to_binary(5) == '101'
assert decimal_to_binary(6) == '110' 
```

<!-- grade_id: htest-d2b, points: 1, grade: true, solution: false -->
```python
# hidden tests
```

### Representing negative numbers

The following also uses 3 bits to represent 8 integers but half of them are negative.

::::{list-table} Encoding positive integers
:name: tab:neg-int
:header-rows: 1

*
    - Binary
    - Decimal
*
    - 000
    - 0
*
    - 001
    - 1
*
    - 010
    - 2
*
    - 011
    - 3
*
    - **1**00
    - **-4**
*
    - **1**01
    - **-3**
*
    - **1**10
    - **-2**
*
    - **1**11
    - **-1**

::::

::::{note}

- The numbers 0, 1, 2, and 3 have the same binary representations as before, but
- the binary representations for 4, 5, 6, and 7 are now used for -4, -3, -2, and -1. 

::::

**What is the benefit of the above representation?**

To subtract a positive binary number $x$ from another positive binary number $y$, i.e.,

$$x - y,$$

it seems we can instead perform the binary addition

$$x + (-y)$$

of a positive binary number $x$ and a negative binary number $-y$.

For example,

$$
\overbrace{011_2}^{3} + \overbrace{100_2}^{-4} = \underbrace{111_2}_{-1}
$$

It seem to work as well if there are bits carried forward, e.g., $1 + (- 3)$ in binary is

$$
\begin{array}{rrrr}
 & 0 & \overset{\color{blue} 1}0 & 1 \\
+ & 1 & 0 & 1 \\\hline
& 1 & 1 &  0 
\end{array}
$$

which translate to $-2$ in decimal as desired.

::::{exercise} overflow
:label: ex:overflow

There is a subtlety when computing $3 - 2$ using binary addition because 

$$ \overbrace{011_2}^3 + \overbrace{110_2}^{-2} = 1001_2$$

which overflows to $4$ bits, a seemingly invalid binary representation. Fortunately, there is a simple fix so binary addition can still apply. Come up with such a fix that also works for other cases such as $3-3$, $2-1$, etc.

:::{hint}
See [two's complement represenation](https://en.wikipedia.org/wiki/Two%27s_complement).

:::

::::

<!-- grade_id: fix-overflow, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Logic Gates

A computer is built from transistors that operates on binary states *on/off*, 
which is also represented as bits *1/0*, or the [*boolean*](https://en.wikipedia.org/wiki/Boolean_data_type#:~:text=In%20computer%20science%2C%20the%20Boolean,of%20logic%20and%20Boolean%20algebra.) value *true/false*. Play with the following simulator to see some examples.

::::{card}
:header: [Logic simulator](https://logic.ly/demo/)

1. Click the logic simulator header above.
2. Select *1-Bit Full Adder*.
3. Change the input and observe the output.

[](#tab:adder) gives the input and output relationship of a simpler adder, called the half adder, which computes the addition of two input bits $A$ and $B$ as

$$
A + B = C\circ S
$$

where $C$ and $S$ are the output bits, and $C\circ S$ is the concatenation of two bits $C$ and $S$.

::::{list-table} Adder.
:name: tab:adder
:header-rows: 1

*
    - A
    - B
    - C$\circ$S
*
    - 0
    - 0
    - 00
*
    - 0
    - 1
    - 01
*
    - 1
    - 0
    - 01
*
    - 1
    - 1
    - 10

::::

The operation can be built using transistors, called [*logic gates*](https://en.wikipedia.org/wiki/Logic_gate), that can carry out basic logic operations such as

- $A \operatorname{AND} B$: which returns $1$ if and only if both $A$ and $B$ are $1$, and $0$ otherwise.
- $A \operatorname{XOR} B$: which returns $1$ if and only if either $A$ and $B$ are $1$ but not both.

The input and output relationships are listed below:

```python
from pandas import DataFrame

DataFrame(
    [[A, B, A & B, A ^ B] for A in (0, 1) for B in (0, 1)],
    columns=["A", "B", "A AND B", "A XOR B"],
)
```

::::{note}

The logical functions are computed using {term}`bitwise operators <bitwise operator>` `&` and `^`, and the result is stored as a `DataFrame` object, which is a useful tool for data analysis to be explained later in the course.

::::

::::{exercise} adder
:label: ex:adder

What is the logic in computing $C$ and $S$ from $A$ and $B$? Use the logic gates AND and XOR only.

::::

<!-- grade_id: adder, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

::::::{seealso}
If you play minecraft, you can build the following devices based on how computers are built from logic gates:

:::::{grid} 1 1 2 2

::::{card} 
:header: An adder for binary addition
:footer: [open in new tab](https://youtu.be/vtpTri-PZXY)

:::{iframe} https://www.youtube.com/embed/vtpTri-PZXY
:width: 100%

:::

::::

::::{card} 
:header: A calculator for calculus
:footer: [open in new tab](https://youtu.be/5NzIt9sns6o)

:::{iframe} https://www.youtube.com/embed/5NzIt9sns6o
:width: 100%

:::

::::

:::::
::::::

## Glossary

:::{glossary}

algorithm
: An algorithm is a set of instructions that a computer program follows to solve a problem or complete a task. It is a step-by-step approach designed to be efficient and accurate.

bitwise operator
: Bitwise operators are used to manipulate individual bits of binary numbers. They allow for the manipulation of data at a low level, such as setting or clearing specific bits, shifting bits left or right, or performing logical operations on bits.

composite data type
: A composite data type groups related data into a single unit so that they can be easily stored and manipulated together.

conditional
: A conditional statement is used to perform different actions based on different conditions. It allows the program to make decisions and execute specific code blocks based on the evaluation of a given condition.

iteration
: An iteration is a repeated execution until a condition is met. It allows for the efficient repetition of tasks without the need for writing the same code multiple times. 

library
: A library is a collection of pre-written code that developers can use to perform common tasks. It contains functions, classes, and other reusable code that can be integrated into a larger software application, saving time and effort.

pseudocode
: Pseudocode is a high-level description of a program or algorithm, using natural language and code-like syntax to express logic without being bound to a specific programming language. It is used during planning and design phases of software development.

:::